Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Runtime$\rightarrow$Factory reset runtime) and then **run all cells** (in the menubar, select Runtime$\rightarrow$Run all).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Fill out your name below in the `NAME` variable, leave the `COLLABORATORS` variable empty.

```python
NAME = "Ritesh Reddy"
COLLABORATORS = ""
```

---

Install required dependency.

```python
!pip install nose
```

Output:
```
Requirement already satisfied: nose in /usr/local/lib/python3.6/dist-packages (1.3.7)
```

Download the required data.

```python
!wget -c 'https://www.dropbox.com/s/senl6bh60a1a93y/HW1_supplement.zip?dl=1' -O data.zip
!unzip -o data.zip
```

Output:
```
--2020-09-17 23:45:58--  https://www.dropbox.com/s/senl6bh60a1a93y/HW1_supplement.zip?dl=1
Resolving www.dropbox.com (www.dropbox.com)... 162.125.1.1, 2620:100:6016:1::a27d:101
Connecting to www.dropbox.com (www.dropbox.com)|162.125.1.1|:443... connected.
HTTP request sent, awaiting response... 301 Moved Permanently
Location: /s/dl/senl6bh60a1a93y/HW1_supplement.zip [following]
--2020-09-17 23:45:58--  https://www.dropbox.com/s/dl/senl6bh60a1a93y/HW1_supplement.zip
Reusing existing connection to www.dropbox.com:443.
HTTP request sent, awaiting response... 302 Found
Location: https://uc5ffd51b27a19fc4a6e5296fb73.dl.dropboxusercontent.com/cd/0/get/A_mxZnm0eSBB6-DorcOFgFb0H2sAkqAD3pt788Naf3XzElMF4OtsFbnt4qFTAKNHv93tUSorj2R-tYVSsb2otz3CURZ6MJvzzW1UOSVA2ixGWqXPXxJtgy_ZRidwGCOM_MU/file?dl=1# [following]
--2020-09-17 23:45:58--  https://uc5ffd51b27a19fc4a6e5296fb73.dl.dropboxusercontent.com/cd/0/get/A_mxZnm0eSBB6-DorcOFgFb0H2sAkqAD3pt788Naf3XzElMF4OtsFbnt4qFTAKNHv93tUSorj2R-tYVSsb2otz3CURZ6MJ
```

Show the downloaded data

- `./BLOSUM62.txt`
- `./Q2b_reference.txt`


```python
!ls *.txt
```

Output:
```
BLOSUM62.txt  Q2b_reference.txt
```

# Sequence Alignment

In bioinformatics, the problem of sequence alignment is reducible to a generalized form of edit distance. Sequence alignment is an instance of the weighted edit distance problem. Here, the cost of the elementary operations may vary according to the symbols involved. This generalization allows us to more accurately model processes in biological sequences that may lead to substitutions, insertions or deletions.

As we saw in lecture, we can store the costs of elementary operations in the form of a scoring matrix $\delta :$ ($\Sigma$ $\cup$ $\{-\})^{2}\rightarrow$ $\mathbb{R}$. The values of this scoring matrix depend on the biological context.

## BLOSUM62

The most widely used scoring matrix in amino acid sequence alignments is known as BLOSUM62. You can read more about BLOSUM here:

* https://www.sciencedirect.com/topics/biochemistry-genetics-and-molecular-biology/blosum
* https://en.wikipedia.org/wiki/BLOSUM

The BLOSUM.txt file that came in the HW1_supplement.zip contains the matrix. You can also download the BLOSUM62 matrix at https://www.ncbi.nlm.nih.gov/Class/FieldGuide/BLOSUM62.txt
 
Note: 
* The `read_blosum62` function is specifically tailored to read only the file above. You should not need to change this function. 
* The matrix uses the `*` character for gaps, but we change it here to the `-` character. 

```python
def read_blosum62(path):
    """
    Reads in the ncbi's BLOSUM62.txt file and loads the scoring matrix
    into a dictionary.
    
    :param: path is the full path in the local filesystem at which the .txt file is located
    :return: a dictionary of dictionaries which will hold the cost of various amino acid
    substitutions as defined in BLOSUM62.
    """
    delta = {}
    with open(path, 'r') as f:
        lines = f.readlines()[6:]
        keys = lines[0].split()
        keys[-1] = '-'
        for i, line in enumerate(lines[1:]):
            delta[keys[i]] = {k : int(v) for (k,v) in zip(keys, line.split()[1:])}  
    return delta

blosum = read_blosum62('./BLOSUM62.txt')
```

## Question 1a - Global Alignment - 10 points

Complete the function `global_align` defined below that uses the Needleman-Wunsch algorithm for global sequence alignment. The function should take two strings, and a scoring function, and compute the score of the optimal global alignment. It should return a tuple `(best score, optimal alignment)`.

We have provided a function `traceback_global` that takes the two strings `v` and `w`, and a list of lists called
`pointers`. This list is what you will use to record backpointers while running your sequence alignment algorithm.
Each `pointers[i][j]` will be one of `UP`, `LEFT`, `TOPLEFT` or `ORIGIN`. The `traceback_global` function will insert gaps into the two strings where necessary, according to the backpointer table you pass it, and return the alignment to you. 

*Hint: Feel free to define helper functions for this and future questions on this notebook.*

<!-- grade_id: cell-dc425c49e7d9bc70, points: 0, grade: false, solution: true -->
```python
UP = (-1,0)
LEFT = (0, -1)
TOPLEFT = (-1, -1)
ORIGIN = (0, 0)

def traceback_global(v, w, pointers):
    i,j = len(v), len(w)
    new_v = []
    new_w = []
    while True:
        di, dj = pointers[i][j]
        if (di,dj) == LEFT:
            new_v.append('-')
            new_w.append(w[j-1])
        elif (di,dj) == UP:
            new_v.append(v[i-1])
            new_w.append('-')
        elif (di,dj) == TOPLEFT:
            new_v.append(v[i-1])
            new_w.append(w[j-1])
        i, j = i + di, j + dj
        if (i <= 0 and j <= 0):
            break
    return ''.join(new_v[::-1])+'\n'+''.join(new_w[::-1])
    

def global_align(v, w, delta):
    """
    Returns the score of the maximum scoring alignment of the strings v and w, as well as the actual alignment as 
    computed by traceback_global. 
    
    :param: v
    :param: w
    :param: delta
    """
    M = [[0 for j in range(len(w)+1)] for i in range(len(v)+1)]
    pointers = [[ORIGIN for j in range(len(w)+1)] for i in range(len(v)+1)]
    score, alignment = None, None
    M[0][0] = 0
    pointers[0][0] = ORIGIN
    for i in range(1,len(v)+1):
      M[i][0] = i*delta[v[i-1]]['-']
      pointers[i][0] = UP

    for j in range(1,len(w)+1):
      M[0][j] = j*delta['-'][w[j-1]]
      pointers[0][j] = LEFT

    for i in range(1,len(v)+1):
      for j in range(1,len(w)+1):
        match_score = M[i-1][j-1] + delta[ v[i-1] ][ w[j-1] ]
        delete_score = M[i-1][j] + delta[ v[i-1] ][ '-' ]
        insert_score = M[i][j-1] + delta[ '-' ][ w[j-1] ]

        M[i][j] = max(match_score, delete_score, insert_score)
        
        if M[i][j] == delete_score:
          pointers[i][j] = UP
        elif M[i][j] == insert_score:
          pointers[i][j] = LEFT
        elif M[i][j] == match_score:
          pointers[i][j] = TOPLEFT     

    # YOUR CODE HERE

    # raise NotImplementedError()
    alignment = traceback_global(v,w, pointers)
    return M[len(v)][len(w)], alignment
```

```python
keys = ['A', 'C', 'T', 'G', '-']
delta = {}
for i in range(len(keys)):
    delta[keys[i]] = {k : v for (k,v) in zip(keys, [1 if keys[i] == keys[j]  else -1 for j in range(len(keys))])}

global_align("TAGATA", "GTAGGCTTAAGGTTA", delta)
```

Output:
```
(-3, '-TAG----A---T-A\nGTAGGCTTAAGGTTA')
```

## Question 1b -- Global Alignment, Protein Data Bank -- 5 points

https://www.rcsb.org/ provides a database of most known proteins. The search functionality allows us to find information on proteins via their *PDB ID* (https://www.rcsb.org/pdb/staticHelp.do?p=help/advancedsearch/pdbIDs.html). You can get the exact amino acid sequence of chains in a protein by searching via the PDB ID, then under the *Display Files* dropdown menu, clicking *FASTA Sequence*.

Use the function `global_align` with the BLOSUM62 scoring matrix to find the alignment score for the following pairs of amino acid sequences
* The A-chain of human insulin (#4F1C:A) and the A-chain of bovine insulin (#2ZP6:A). Set the variable `scoreA` below to this value.
* The B-chain of the above proteins. Set the variable `scoreB` below to this value.
* Set the `alignmentA` and `alignmentB` variables to the output of `traceback_global` function for both chains.

<!-- grade_id: cell-41162815d1c04b75, points: 0, grade: false, solution: true -->
```python
scoreA = None
scoreB = None
alignmentA = None
alignmentB = None
scoreA = global_align('GIVEQCCTSICSLYQLENYCN','GIVEQCCASVCSLYQLENYCN',blosum)[0]
scoreB = global_align('FVNQHLCGSHLVEALYLVCGERGFFYTPKT','FVNQHLCGSHLVEALYLVCGERGFFYTPKA',blosum)[0]
alignmentA = global_align('GIVEQCCTSICSLYQLENYCN','GIVEQCCASVCSLYQLENYCN',blosum)[1]
alignmentB = global_align('FVNQHLCGSHLVEALYLVCGERGFFYTPKT','FVNQHLCGSHLVEALYLVCGERGFFYTPKA',blosum)[1]
# print(alignmentA, scoreA,'\n', alignmentB, scoreB)

# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-49569ad025096c67, points: 15, grade: true, solution: false -->
```python
import nose.tools as Test

# Student test cases (You may add more)
Test.assert_equal(global_align('LRRGEPVY', 'LEKGDTLYILVG', blosum), (5, 'LRRGEPVY----\nLEKGDTLYILVG'))
Test.assert_equal(global_align('ALYFFT', 'QCASVT', blosum), (-2, 'ALYFFT\nQCASVT'))
keys = ['A', 'C', 'T', 'G', '-']
delta = {}
for i in range(len(keys)):
    delta[keys[i]] = {k : v for (k,v) in zip(keys, [1 if keys[i] == keys[j]  else -1 for j in range(len(keys))])}
Test.assert_equal(global_align('GTTTTATAGC', 'TGGATCGAT', delta),(-3, 'GTTTTATAGC-\n-TGG-ATCGAT'))

```

## Question 2a - Fitting Alignment - 10 points

As we know, next generation sequencing doesn't produce a fully sequenced genome, but rather many short reads. Here, we use fitting alignment. 

Complete the `fitting_align` function defined below. The function should take a short read string, a reference genome string, and a scoring matrix, and return the score of the maximum scoring alignment of the short read and all substrings of the reference string, along with the actual fitting alignmnet as computed by the given function `traceback_fitting`. This function is similar to `traceback_global` but also requires you to pass in your dynamic programming table `M`, as well as the starting column for the backtrace `init_j`. 

** Note: In the case of ties with respect to determining the starting column for the backtrace, use the rightmost.**

<!-- grade_id: cell-5907258be80654a4, points: 0, grade: false, solution: true -->
```python
UP = (-1,0)
LEFT = (0, -1)
TOPLEFT = (-1, -1)
ORIGIN = (0, 0)

def traceback_fitting(v, w, M, init_j, pointers):
    i, j = len(v), init_j
    new_v = []
    new_w = []
    while True:
        di, dj = pointers[i][j]
        if (di,dj) == LEFT:
            new_v.append('-')
            new_w.append(w[j-1])
        elif (di,dj) == UP:
            new_v.append(v[i-1])
            new_w.append('-')
        elif (di,dj) == TOPLEFT:
            new_v.append(v[i-1])
            new_w.append(w[j-1])
        i, j = i + di, j + dj
        if (i <= 0):
            break
    return ''.join(new_v[::-1]) + '\n'+''.join(new_w[::-1])
    
def fitting_align(short, reference, delta):
    """
    Returns the score of the maximum scoring alignment of short and all 
    substrings of reference. 
    
    :param: short the shorter of the two strings we are trying to align    
    :param: reference the longer string among whose substrings we are doing global alignment
    :param: delta the scoring function for the alphabet of the two strings
    
    :returns: a tuple (score, alignment)
    """
    M = [[0 for j in range(len(reference)+1)] for i in range(len(short)+1)]     
    pointers = [[ORIGIN for j in range(len(reference)+1)] for i in range(len(short)+1)] 
    score = None
    init_j = 0
    v = short
    w = reference
    # YOUR CODE HERE
    for i in range(1,len(v)+1):
      M[i][0] = i*delta[v[i-1]]['-']
      pointers[i][0] = UP

    for j in range(1,len(w)+1):
      M[0][j] = 0
      pointers[0][j] = ORIGIN

    
    for i in range(1,len(v)+1):
      for j in range(1,len(w)+1):
        match_score = M[i-1][j-1] + delta[ v[i-1] ][ w[j-1] ]
        delete_score = M[i-1][j] + delta[ v[i-1] ][ '-' ]
        insert_score = M[i][j-1] + delta[ '-' ][ w[j-1] ]

        M[i][j] = max(match_score, delete_score, insert_score)

        if M[i][j] == delete_score:
          pointers[i][j] = UP
        elif M[i][j] == insert_score:
          pointers[i][j] = LEFT
        elif M[i][j] == match_score:
          pointers[i][j] = TOPLEFT 

    max_pos = 0
    max_score = 0
    for j in range(1,len(w)+1):
      if M[len(v)][j] > max_score:
        max_score = M[i][j]
        max_pos = j

    init_j = max_pos

    # raise NotImplementedError()
    alignment = traceback_fitting(short,reference,M, init_j,pointers)
    return M[len(v)][init_j], alignment
```

```python
keys = ['A', 'C', 'T', 'G', '-']
delta = {}
for i in range(len(keys)):
    delta[keys[i]] = {k : v for (k,v) in zip(keys, [1 if keys[i] == keys[j]  else -1 for j in range(len(keys))])}
    
fitting_align("TAGATA", "GTAGGCTTAAGGTTA", delta)
# print(delta)
```

Output:
```
(2, 'TAGA-T-A\nTAGGCTTA')
```

<!-- grade_id: cell-057e4f4d500c3780, points: 10, grade: true, solution: false -->
```python
import nose.tools as Test

# Define scoring matrix
keys = ['A', 'C', 'T', 'G', '-']
delta = {}
for i in range(len(keys)):
    delta[keys[i]] = {k : v for (k,v) in zip(keys, [1 if keys[i] == keys[j]  else -1 for j in range(len(keys))])}
delta['-']['-'] = -1

# Student test cases (You can add more)
Test.assert_equal(fitting_align('ACTGCT', 'ACTGTCGTACGTGTACGTGCTATTACGATTCGGATGCATTGTGCATTTGGGCGATCTTATTCTTATC', delta),\
                  (5, 'AC-TGCT\nACGTGCT'))
Test.assert_equal(fitting_align('PPPVP','MHHHHHHSSPIDPPGKPVPLNITRHTVTLKWAKPEYTGGFKITSYIVEKRDLPNGRWLKANFSNILENEFTVSGLTEDAA\
YEFRVIAKNAAGAISPPSEPSDAITCRDDVEA', blosum), (24, 'PP--PVP\nPPGKPVP'))

```

## Question 2b - Fitting Alignment - 5 points

After writing the function use it to align the reference sequence `./Q2b_reference.txt` and the short read provided below. The reference is a partial version of the human Y-chromosome reference found at ftp://ftp.ncbi.nih.gov/genomes/H_sapiens/CHR_Y/. 

Set the variable `score_fitting` to the fitting alignment score, and set the `alignment_fitting` to the output of the `traceback_fitting` function. You can use the scoring matrix defined below. 

** Note: The reference sequence is very large, and so the fitting alignment will take a few minutes. Make sure you use dynamic programming to implement your alignment functions, or the cell will never finish. **

<!-- grade_id: cell-4e3457bc42a54b0d, points: 5, grade: true, solution: false -->
```python
short = 'AGAAGTGAGTTTTGGATAGTAAAATAAGTTTCGAACTCTGGCACCTTTCAATTTTGTCGCACTCTCCTTGTTTTTGACAATGCAATCATATGCTTCTGCT'
keys = ['A', 'C', 'T', 'G', '-']
delta_fitting = {}
for i in range(len(keys)):
    delta_fitting[keys[i]] = {k : v for (k,v) in zip(keys, [1 if keys[i] == keys[j]  else -1 for j in range(len(keys))])}
delta_fitting['-']['-'] = -1

f = open("./Q2b_reference.txt", "r")
# Read the reference sequence, and call your fittinga alignment function
score_fitting = None
alignment_fitting = None
score_fitting = fitting_align(short, f.read(), delta_fitting)[0]
alignment_fitting = fitting_align(short, f.read(), delta_fitting)[1]
print(score_fitting)
```

Output:
```
33
```

## Question 2c - 5 points

Find the name of the gene the short read in the question above originated from. Save your answer in variable `ANSWER` below. 

*Hint : There is an online tool that will do this for you*

<!-- grade_id: cell-39bf2f4158cb333b, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
ANSWER = 'KRTAP5-5'
# raise NotImplementedError()
```

## Question 3 - Local Alignment - 15 points

Finally, we use local alignment for things like aligning functional units of proteins called domains. 

Use the Smith-Waterman algorithm to complete the function `local_align` defined below that takes two strings `v` and `w` and a scoring matrix, and returns the maximum scoring global alignment across all substrings of `v` and `w`, as well as the alignment string produced by `traceback_local`. As with the above functions, you must record backpointers. This function is very similar to `traceback_fitting`, except it also takes the starting cell for the back trace in the form of coordinates `init_i, init_j`. 

** Note: In the case of ties with respect to determining the starting column for the backtrace, use the bottom-right most cell.**

<!-- grade_id: cell-0f8d6da9a1e58ec1, points: 0, grade: false, solution: true -->
```python
UP = (-1,0)
LEFT = (0, -1)
TOPLEFT = (-1, -1)
ORIGIN = (0, 0)

def traceback_local(v, w, M, init_i, init_j, pointers):
    i,j = init_i, init_j
    new_v = []
    new_w = []
    while True:
        di, dj = pointers[i][j]
        if (di,dj) == LEFT:
            new_v.append('-')
            new_w.append(w[j-1])
        elif (di,dj) == UP:
            new_v.append(v[i-1])
            new_w.append('-')
        elif (di,dj) == TOPLEFT:
            new_v.append(v[i-1])
            new_w.append(w[j-1])
        i, j = i + di, j + dj
        if (M[i][j] == 0):
            break
    return ''.join(new_v[::-1]) + '\n'+''.join(new_w[::-1])


def local_align(v, w, delta):
    """
    Returns the score of the maximum scoring alignment of all possible substrings of v and w. 
    
    :param: v
    :param: w
    :param: delta
    """
    M = [[0 for j in range(len(w)+1)] for i in range(len(v)+1)]        
    pointers = [[ORIGIN for j in range(len(w)+1)] for i in range(len(v)+1)]
    # M[0][0] = -1
    for i in range(1,len(v)+1):
      M[i][0] = 0
      pointers[i][0] = ORIGIN

    for j in range(1,len(w)+1):
      M[0][j] = 0
      pointers[0][j] = ORIGIN

    for i in range(1,len(v)+1):
      for j in range(1,len(w)+1):
        match_score = M[i-1][j-1] + delta[ v[i-1] ][ w[j-1] ]
        delete_score = M[i-1][j] + delta[ v[i-1] ][ '-' ]
        insert_score = M[i][j-1] + delta[ '-' ][ w[j-1] ]

        M[i][j] = max(0,match_score, delete_score, insert_score)

        if M[i][j] == delete_score:
          pointers[i][j] = UP
        elif M[i][j] == insert_score:
          pointers[i][j] = LEFT
        elif M[i][j] == match_score:
          pointers[i][j] = TOPLEFT
        elif M[i][j] == 0:  
          pointers[i][j] = ORIGIN

    max_pos = None
    max_score = 0
    for i in range(1, len(v)+1):
        for j in range(1, len(w)+1):
            if M[i][j] >= max_score:
                max_score = M[i][j]
                max_pos   = (i, j)
    init_i, init_j = max_pos

    # YOUR CODE HERE
    # raise NotImplementedError()
    alignment = traceback_local(v, w, M, init_i, init_j , pointers)
    return M[init_i][init_j], alignment 
```

```python
keys = ['A', 'C', 'T', 'G', '-']
delta = {}
for i in range(len(keys)):
    delta[keys[i]] = {k : v for (k,v) in zip(keys, [1 if keys[i] == keys[j]  else -1 for j in range(len(keys))])}
    
local_align("TAGATA", "GTAGGCTTAAGGTTA", delta)
```

Output:
```
(3, 'TAG\nTAG')
```

Use your `local_align` function to find the common domain between the two amino acid sequences given below. 
* Set the `score_local` variable to the output of your function. 
* Set the `alignment_local` variable to the output of the `traceback_local` function. 

<!-- grade_id: cell-f37ba5b055e2d307, points: 0, grade: false, solution: true -->
```python
LTK_MOUSE = "MGCSHRLLLWLGAAGTILCSNSEFQTPFLTPSLLPVLVLNSQEQKVTPTP\
SKLEPASLPNPLGTRGPWVFNTCGASGRSGPTQTQCDGAYTGSSVMVTVG\
AAGPLKGVQLWRVPDTGQYLISAYGAAGGKGAQNHLSRAHGIFLSAVFFL\
RRGEPVYILVGQQGQDACPGGSPESQLVCLGESGEHATTYGTERIPGWRR\
WAGGGGGGGGATSIFRLRAGEPEPLLVAAGGGGRSYRRRPDRGRTQAVPE\
RLETRAAAPGSGGRGGAAGGGSGWTSRAHSPQAGRSPREGAEGGEGCAEA\
WAALRWAAAGGFGGGGGACAAGGGGGGYRGGDTSESDLLWADGEDGTSFV\
HPSGELYLQPLAVTEGHGEVEIRKHPNCSHCPFKDCQWQAELWTAECTCP\
EGTELAVDNVTCMDLPTTASPLILMGAVVAALALSLLMMCAVLILVNQKC\
QGLWGTRLPGPELELSKLRSSAIRTAPNPYYCQVGLSPAQPWPLPPGLTE\
VSPANVTLLRALGHGAFGEVYEGLVTGLPGDSSPLPVAIKTLPELCSHQD\
ELDFLMEALIISKFSHQNIVRCVGLSFRSAPRLILLELMSGGDMKSFLRH\
SRPHPGQLAPLTMQDLLQLAQDIAQGCHYLEENHFIHRDIAARNCLLSCS\
GASRVAKIGDFGMARDIYQASYYRKGGRTLLPVKWMPPEALLEGLFTSKT\
DSWSFGVLLWEIFSLGYMPYPGHTNQEVLDFIATGNRMDPPRNCPGPVYR\
IMTQCWQHQPELRPDFGSILERIQYCTQDPDVLNSPLPVEPGPILEEEEA\
SRLGNRSLEGLRSPKPLELSSQNLKSWGGGLLGSWLPSGLKTLKPRCLQP\
QNIWNPTYGSWTPRGPQGEDTGIEHCNGSSSSSIPGIQ"

ATK_MOUSE = "MGAAGFLWLLPPLLLAAASYSGAATDQRAGSPASGPPLQPREPLSYSRLQ\
RKSLAVDFVVPSLFRVYARDLLLPQPRSPSEPEAGGLEARGSLALDCEPL\
LRLLGPLPGISWADGASSPSPEAGPTLSRVLKGGSVRKLRRAKQLVLELG\
EETILEGCIGPPEEVAAVGILQFNLSELFSWWILHGEGRLRIRLMPEKKA\
SEVGREGRLSSAIRASQPRLLFQIFGTGHSSMESPSETPSPPGTFMWNLT\
WTMKDSFPFLSHRSRYGLECSFDFPCELEYSPPLHNHGNQSWSWRHVPSE\
EASRMNLLDGPEAEHSQEMPRGSFLLLNTSADSKHTILSPWMRSSSDHCT\
LAVSVHRHLQPSGRYVAQLLPHNEAGREILLVPTPGKHGWTVLQGRVGRP\
ANPFRVALEYISSGNRSLSAVDFFALKNCSEGTSPGSKMALQSSFTCWNG\
TVLQLGQACDFHQDCAQGEDEGQLCSKLPAGFYCNFENGFCGWTQSPLSP\
HMPRWQVRTLRDAHSQGHQGRALLLSTTDILASEGATVTSATFPAPMKNS\
PCELRMSWLIRGVLRGNVSLVLVENKTGKEQSRTVWHVATDEGLSLWQHT\
VLSLLDVTDRFWLQIVTWWGPGSRATVGFDNISISLDCYLTISGEEKMSL\
NSVPKSRNLFEKNPNKESKSWANISGPTPIFDPTVHWLFTTCGASGPHGP\
TQAQCNNAYQNSNLSVVVGSEGPLKGVQIWKVPATDTYSISGYGAAGGKG\
GKNTMMRSHGVSVLGIFNLEKGDTLYILVGQQGEDACPRANQLIQKVCVG\
ENNVIEEEIRVNRSVHEWAGGGGGGGGATYVFKMKDGVPVPLIIAAGGGG\
RAYGAKTETFHPERLESNSSVLGLNGNSGAAGGGGGWNDNTSLLWAGKSL\
LEGAAGGHSCPQAMKKWGWETRGGFGGGGGGCSSGGGGGGYIGGNAASNN\
DPEMDGEDGVSFISPLGILYTPALKVMEGHGEVNIKHYLNCSHCEVDECH\
MDPESHKVICFCDHGTVLADDGVSCIVSPTPEPHLPLSLILSVVTSALVA\
ALVLAFSGIMIVYRRKHQELQAMQMELQSPEYKLSKLRTSTIMTDYNPNY\
CFAGKTSSISDLKEVPRKNITLIRGLGHGAFGEVYEGQVSGMPNDPSPLQ\
VAVKTLPEVCSEQDELDFLMEALIISKFNHQNIVRCIGVSLQALPRFILL\
ELMAGGDLKSFLRETRPRPNQPTSLAMLDLLHVARDIACGCQYLEENHFI\
HRDIAARNCLLTCPGAGRIAKIGDFGMARDIYRASYYRKGGCAMLPVKWM\
PPEAFMEGIFTSKTDTWSFGVLLWEIFSLGYMPYPSKSNQEVLEFVTSGG\
RMDPPKNCPGPVYRIMTQCWQHQPEDRPNFAIILERIEYCTQDPDVINTA\
LPIEYGPVVEEEEKVPMRPKDPEGMPPLLVSPQPAKHEEASAAPQPAALT\
APGPSVKKPPGAGAGAGAGAGAGPVPRGAADRGHVNMAFSQPNPPPELHK\
GPGSRNKPTSLWNPTYGSWFTEKPAKKTHPPPGAEPQARAGAAEGGWTGP\
GAGPRRAEAALLLEPSALSATMKEVPLFRLRHFPCGNVNYGYQQQGLPLE\
ATAAPGDTMLKSKNKVTQPGP"

score_local, alignment_local = None, None
score_local = local_align(LTK_MOUSE, ATK_MOUSE, blosum)[0]
alignment_local = local_align(LTK_MOUSE, ATK_MOUSE, blosum)[1]
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-5b737423ad2b4fc0, points: 15, grade: true, solution: false -->
```python
import nose.tools as Test
# Student test cases (You may add more)
keys = ['A', 'C', 'T', 'G', '-']
delta_local = {}
for i in range(len(keys)):
    delta_local[keys[i]] = {k : v for (k,v) in zip(keys, [1 if keys[i] == keys[j]  else -1 for j in range(len(keys))])}
delta_local['-']['-'] = -1
# Test.assert_equal(local_align("TACGTGACGTCTATCAT", "TTGTGCATGTATCTGAC", delta_local), (6, 'GTGACGTCTATCAT\nGTG-CATGTATC-T'))
Test.assert_equal(local_align("ACTGATTTCGATGCTGTACGTGACGTACGTTTATTCTATCAT", "TCTGACTGTGACTATACTATTGTGCATGTATCTGACTAGCTAG", delta_local),\
                  (13, 'ACTGATTTCGATGC-TGTA-CGTGACGTACGTTTA\nACT-ATT--G-TGCATGTATC-TGAC-TA-GCT-A'))

```